#include <stdio.h>

// 接下来, 让我们来学习更多整数的表示方式
// 用 0b 开头表示二进制数 (b不区分大小写)
// 用 0  开头表示八进制数
// 用 0x 开头表示十六进制数 (不区分大小写)

// I AM NOT DONE

int main() {

    int binary_lowercase = 0b101;
    int binary_uppercase = 0B101;

    int octonary = 017;

    int hexadecimal_uppercase = 0X1A;
    int hexadecimal_lowercase = 0x1a;
    int hexadecimal_mixed = 0X1a;

    printf("binary_lowercase = %d\n", binary_lowercase);
    printf("binary_uppercase = %d\n", binary_uppercase);

    printf("octonary = %d\n", octonary);

    printf("hexadecimal_uppercase = %d\n", hexadecimal_uppercase);
    printf("hexadecimal_lowercase = %d\n", hexadecimal_lowercase);
    printf("hexadecimal_mixed = %d\n", hexadecimal_mixed);

    // 请根据上面的示例代码, 声明你的变量, 补充下面的空白


    printf("/*  TODO  */ = %d\n", /*  TODO  */);
    printf("/*  TODO  */ = %d\n", /*  TODO  */);

    printf("/*  TODO  */ = %d\n", /*  TODO  */);

    printf("/*  TODO  */ = %d\n", /*  TODO  */);
    printf("/*  TODO  */ = %d\n", /*  TODO  */);
    printf("/*  TODO  */ = %d\n", /*  TODO  */);

    return 0;
}