// 知识点:  一、C程序的入口是 main 函数
//             就像一栋房子不能没有门一样, 一个C程序也不能没有 main 函数
//         二、C语言的每一条语句都必须以  英文  的分号 ; 结尾
// 下面

// I AM NOT DONE

int mian() {
    // 什么都没做
    return 0；
}

// 补充的知识, 好奇的同学可以看看
// 一个C程序的入口就是 main 函数(至于函数是什么，你现在可以简单地理解为一系列操作的集合)
// 运行一个C程序, 就会从 main 函数 进入, 再按顺序一条一条地执行语句

// int 表示该函数的返回值, 不过现在你不用在意
// () 表示该函数的参数列表, 显然这个函数没有参数, 不过现在你不用在意
// return 就像它中文意思一样: 返回; 不过既然可以没参数, 那么同样地, 一个函数也可以没有返回值