// 在进行练习之前请打开位于 exercise 文件夹中的名为 C语言基础概念 的文件
// 里面的内容是对C语言的基础概念的总述

// 知识点:  C程序的入口是 main 函数 (至于函数是什么，你现在可以简单地理解为一系列操作的集合)
//          就像一栋房子不能没有门一样, 一个C程序也不能没有 main 函数

// 练习: 参照知识点, 改正下面的代码, 保存文件, 看一下终端的输出信息, 直到通过编译
// 在终端输入 hint 获取提示 

// I AM NOT DONE

int mian(void) {
    // 什么都没做
    return 0;
}

// 补充的知识
// 运行一个C程序, 就会从 main 函数 进入, 再按顺序一条一条地执行语句