// 知识点:  一、C程序的入口是 main 函数 (至于函数是什么，你现在可以简单地理解为一系列操作的集合)
//             就像一栋房子不能没有门一样, 一个C程序也不能没有 main 函数
//         二、C语言的每一条语句都必须以  英文  的分号 ; 结尾

// 练习: 参照知识点, 改正下面的代码, 保存文件, 看一下终端的输出信息, 直到通过编译
// 在终端输入 hint 获取提示 

// I AM NOT DONE

int mian() {
    // 什么都没做
    return 0；
}

// 补充的知识
// 运行一个C程序, 就会从 main 函数 进入, 再按顺序一条一条地执行语句